---
title: 'Out of sample prediction'
description: 'Out of sample prediction with Python'
---

## Introduction

```yaml
type: NormalExercise
key: dbc917f41c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
